---
title       : Your first steps in Looker
description : Welcome!
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf
  


--- type:VideoExercise lang:python xp:50 skills:2 key:0d52355f10
## Doing stuff in Looker!


*** =video_link
//player.vimeo.com/video/154783078
--- type:PlainMultipleChoiceExercise lang:python xp:50 skills:1 key:4e613a5cf1
## How many puppies ? 
![](https://i.ytimg.com/vi/Fu2ztMCfCQM/hqdefault.jpg)

You would like to know how many orders we have delivered in the zone MILA. 
Please navigate to the correct explore and answer below:

*** =instructions
- 50
- 100
- 500
- 1500

*** =hint
Have we delivered all orders?


*** =sct
```{python}
# SCT written with pythonwhat: https://github.com/datacamp/pythonwhat/wiki

msg_bad = "That is not correct!"
msg_success = "That's correct!"
test_mc(4, [msg_bad, msg_bad, msg_bad, msg_success])
```



--- type:VideoExercise lang:python xp:50 skills:2 key:99ddf810bc
## Next Section


*** =video_link
//player.vimeo.com/video/154783078


